```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Lab 02: The `LETTERS` string

## Today's lab

The goals of this lab activity:
* Understand the `LETTERS` string, and how it can convert characters into a numerical representation
* Perform modular reduction using Python's `%` operator
* Encipher individual letters using the Caesar shift

This lab covers parts of [Chapter 4](https://macs4200.org/chapters/04/2/caesar-cipher) of the online textbook. You should read the examples in the book, but not right now. Instead, let's get started!

<hr style="border: 5px solid #346094;" />
<hr style="border: 1px solid #f4c300;" />

# Part 1: Our Alphabet

In this course, unless otherwise noted, you can assume plaintext and ciphertext are characters from the English language. We will need a reliable and consistent way of converting these characters into an equivalent numerical form.

While there are more sophisticated and modern ways of converting characters to numbers, and vice versa, in this course we'll use a string called `LETTERS` to hold the characters of our alphabet, where the index of the character is the numerical representation of that character. This allows us to use string indexing, `LETTERS[...]` to quickly compute the characters that correspond to a particular numerical value.

```python
LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ'
```

For example, 
*  Since `LETTERS[0]` is `'A'`, $A \rightarrow 0$
*  Since `LETTERS[3]` is `'D'`, $D \rightarrow 3$

So, the characters `M`, `A`, `T`, and `H` could be converted to the equivalent numerical values of: `12`, `0`, `19`, `7`

### Question 1

Suppose we "extend" our usual alphabet to include lowercase letters and some punctuation. We'll call it `BIGLETTERS` and define it in the code cell below.

```python
BIGLETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz!@#$%^&*()'
```

Fill in the `...` in the code cell below to compute the characters in `BIGLETTERS` that correspond to the numerical values of 31, 44, 47, and 53. You should use string indexing to determine the values.

```python
character_31 = ...
character_44 = ...
character_47 = ...
character_53 = ...

print(character_31, character_44, character_47, character_53)
```

```python
grader.check("q1")
```

### Question 2

We could even create an "emoji alphabet", where the alphabet contains every emoji. We'll call this alphabet `EMOJI`.

```python
EMOJI = '😀😃😄😁😆🥹😅😂🤣😍😌😉🙃🙂😇😊☺️🥲🥰😘😗😙😚😋😛😝😜😏🥳🤩🥸😎🤓🧐🤨🤪😒😞😔😟😕🙁☹️😣😖🤬😡😠😤😭😢🥺😩😫🤯😳🥵🥶😶‍🌫️😱😨😰😥🫠🤫🫡🫢🤭🫣🤔🤗😓🤥😶🫥😐🫤😑😬🙄😯😮‍💨😪🤤😴🥱😲😮😧😦😵😵‍💫🤐🥴🤢🤮🤧😷🤒💩🤡👺👹👿😈🤠🤑🤕👻💀☠️👽👾🤖🎃😺😸🤲🫶😾😿🙀😽😼😻😹👐🙌👏🤝👍👎✊🤛👊🤏🤌👌🤘🤟🫰✌️🤞🤜🫳🫴👈👉👆👇☝️✋🤚🖕🦾💪🫱🫲🤙👋🖖🖐✍️🙏🫵🦶🦵🦿💄💋👄👀👁👣👃🦻👂👅🦷🫦🫀🫁🧠🗣👤👥🫂👶👧👩‍🦰👨‍🦱🧑‍🦱👩‍🦱👨🧑👩👦🧒🧑‍🦰👨‍🦰👱‍♀️👱👱‍♂️👩‍🦳🧑‍🦳👨‍🦳👩‍🦲👲👴🧓👵🧔‍♂️🧔🧔‍♀️👨‍🦲🧑‍🦲👳‍♀️👳👳‍♂️🧕👮‍♀️👮👮‍♂️👷‍♀️👷🧑‍⚕️👩‍⚕️🕵️‍♂️🕵️🕵️‍♀️💂‍♂️💂💂‍♀️👷‍♂️👨‍⚕️👩‍🌾🧑‍🌾👩‍🍳🧑‍🍳👨‍🍳👩‍🎓🧑‍🎓👨‍🌾🧑‍🏭👩‍🏭👨‍🏫🧑‍🏫👩‍🏫👨‍🎤🧑‍🎤👩‍🎤👨‍🎓👨‍🏭👩‍💻🧑‍💻👨‍💻👩‍💼🧑‍💼👨‍💼👩‍🔧🧑‍🔧🧑‍🚒👩‍🚒👨‍🎨🧑‍🎨👩‍🎨👨‍🔬🧑‍🔬👩‍🔬👨‍🔧👨‍🚒👩‍✈️🧑‍✈️👨‍✈️👩‍🚀🧑‍🚀👨‍🚀👩‍⚖️🧑‍⚖️🫅👸🤵‍♂️🤵🤵‍♀️👰‍♂️👰👰‍♀️👨‍⚖️🤴🥷🦸‍♀️🦸🦸‍♂️🦹🦹‍♂️🤶🦹‍♀️🧌🧝‍♂️🧝🧝‍♀️🧙‍♂️🧙🧙‍♀️🎅🧑‍🎄🧛‍♀️🧛🧛‍♂️🧟‍♀️🧟🧟‍♂️🧞‍♀️🧞🧞‍♂️🫄🤰👼🧚‍♂️🧚🧚‍♀️🧜‍♂️🧜🧜‍♀️🫃🤱👩‍🍼🧑‍🍼👨‍🍼🙇‍♀️🙇🙇‍♂️💁‍♀️🙋‍♀️🙆‍♂️🙆🙆‍♀️🙅‍♂️🙅🙅‍♀️💁‍♂️💁🙋🙋‍♂️🧏‍♀️🧏🧏‍♂️🤦‍♀️🤦🤦‍♂️🤷‍♀️💇‍♀️🙍‍♂️🙍🙍‍♀️🙎‍♂️🙎🙎‍♀️🤷‍♂️🤷💇💇‍♂️💆‍♀️💆💆‍♂️🧖‍♀️🧖🧖‍♂️💅🧑‍🦽👩‍🦽🕴👯‍♂️👯👯‍♀️🕺💃🤳👨‍🦽🧑‍🦼👨‍🦼🚶‍♀️🚶🚶‍♂️👩‍🦯🧑‍🦯🧍🧍‍♀️🏃‍♂️🏃🏃‍♀️🧎‍♂️🧎🧎‍♀️👨‍🦯🧍‍♂️👫👭👬👩‍❤️‍👨👩‍❤️‍👩💑👨‍❤️‍👨👩‍❤️‍💋‍👨👩‍👩‍👦👨‍👩‍👧‍👧👨‍👩‍👦‍👦👨‍👩‍👧‍👦👨‍👩‍👧👨‍👩‍👦👨‍❤️‍💋‍👨💏👩‍❤️‍💋‍👩👩‍👩‍👧👩‍👩‍👧‍👦👩‍👩‍👦‍👦👩‍👩‍👧‍👧👨‍👨‍👦👨‍👨‍👧👨‍👨‍👧‍👦👨‍👨‍👦‍👦👨‍👨‍👧‍👧👨‍👦‍👦👨‍👧‍👦👨‍👧👨‍👦👩‍👧‍👧👩‍👦‍👦👩‍👧‍👦👩‍👧👩‍👦👨‍👧‍👧🪢🧶🧵🪡🧥🥼🦺👚👘🩱👙👗👔🩳🩲👖👕🥻🩴🥿👠👡👢👞👟🥾🪖⛑👒🎓🧢🎩🧣🧤🧦👑💍👝👛👜💼🎒🧳👓🌂🥽🕶❤️🧡💛💚💙💜🤎🖤🤍💖💗💓💞💕❣️❤️‍🩹❤️‍🔥💔💘💝♡♥︎❥❦❧☙❢এლდღɞʚෆᰔᩚᰔঞଓᜊᥫ᭡ꨄஐᦗ🐶🐱🐭🐹🐰🦊🐻🐼🐻‍❄️🙈🐵🐸🐽🐷🐮🦁🐯🐨🙉🙊🐒🐔🐧🐦🐤🐣🐥🐝🦄🐴🐗🐺🦇🦉🦅🦆🪱🐛🦋🐌🐞🐜🪰🪲🪳🦖🦎🐍🐢🦂🕸🕷🦗🦟🦕🐙🦑🦐🦞🦀🐡🐠🐟🦓🐆🐅🐊🦭🦈🐋🐳🐬🦍🦧🦣🐘🦛🦏🐪🐫🦒🐑🐏🐖🐎🐄🐂🐃🦬🦘🦙🐐🦌🐕🐩🦮🐕‍🦺🐈🐈‍⬛🕊🦩🦢🦜🦚🦤🦃🐓🪶🐇🦝🦨🦡🦫🦦🦥🐁🐀🌳🌲🎄🌵🐲🐉🐾🦔🐿🌴🪵🌱🌿☘️🍀🎍🪴🎋🪨🪸🐚🍄🪹🪺🍁🍂🍃🌾💐🌷🌹🥀🪷🌺🌸🌼🌗🌖🌕🌚🌜🌛🌝🌞🌻🌘🌑🌒🌓🌔🌙🌎🌍🌏🔥💥☄️⚡️✨🌟⭐️💫🪐🌪🌈☀️🌤⛅️🌥☁️🌦🌧💧💨🌬⛄️☃️❄️🌨🌩⛈💦🫧☔️☂️🌊🍏🍎🍐🍊🍋🍌🍉🍇🍓🍅🥝🥥🍍🥭🍑🍒🍈🫐🍆🥑🥦🥬🥒🌶🫑🌽🥕🥖🍞🥯🥐🍠🥔🧅🧄🫒🥨🧀🥚🍳🧈🥞🧇🥓🥩🥪🫓🍕🍟🍔🌭🦴🍖🍗🥙🧆🌮🌯🫔🥗🥘🫕🥫🦪🥟🍱🍣🍛🍲🍜🍝🫙🍤🍙🍚🍘🍥🥠🥮🍢🍡🍭🍮🎂🍰🧁🥧🍦🍨🍧🍬🍫🍿🍩🍪🌰🥜🫘🍯🧋🥤🧃🍵☕️🫖🍼🫗🥛🍶🍺🍻🥂🍷🥃🍸🍹🧉🧂🥢🥡🥣🍽🍴🥄🧊🍾⚽️🏀🏈⚾️🥎🎾🏐🏉🥏🪃🏏🥍🏑🏒🏸🏓🪀🎱🥅⛳️🪁🛝🏹🎣🤿🥊🥋🏂⛷🎿🥌⛸🛷🛼🛹🎽🪂🏋️‍♀️🏋️🏋️‍♂️🤼‍♀️🤼🤼‍♂️🤸‍♀️🤸🏌️‍♀️🤾‍♂️🤾🤾‍♀️🤺⛹️‍♂️⛹️⛹️‍♀️🤸‍♂️🏌️🏌️‍♂️🏇🧘‍♀️🧘🧘‍♂️🏄‍♀️🏄🏄‍♂️🚣‍♂️🚣🚣‍♀️🤽‍♂️🤽🤽‍♀️🏊‍♂️🏊🏊‍♀️🧗‍♀️🧗🧗‍♂️🚵‍♀️🚵🚵‍♂️🚴‍♀️🚴🚴‍♂️🎫🎗🏵🎖🏅🥉🥈🥇🏆🎟🎪🤹‍♀️🤹🤹‍♂️🎭🩰🎨🎬🪗🎺🎷🪘🥁🎹🎼🎧🎤🎸🪕🎻🎲♟🎯🎳🎮🎰🧩🚗🚕🚙🚌🚎🏎🚓🚑🚒🩼🦼🦽🦯🚜🚛🚚🛻🚐🛴🚲🛵🏍🛺🛞🚨🚔🚍🚝🚞🚋🚃🚟🚠🚡🚖🚘🚄🚅🚈🚂🚆🚇🚊🚉✈️🛶🚁🛸🚀🛰💺🛩🛬🛫⛵️🚤🛥🛳🚢🛟⛴⚓️🪝🗼🗽🗿🗺🚏🚥🚦🚧⛽️🏰🏯🏟🎡🎢🎠⛲️⛱🏖🛖⛺️🏕🗻🏔⛰🌋🏜🏝🏠🏡🏘🏚🏗🏭🏢🏬🏣🏛💒🏩🏫🏪🏨🏦🏥🏤⛪️🕌🕍🛕🕋⛩🛤🛣🗾🌆🌇🎆🎇🌠🌄🌅🏞🎑🏙🌃🌌🌉🌁⌚️📱📲💻⌨️🖥🖨🖱🖲📸📷📼📀💿💾💽🗜🕹📹🎥📽🎞📞☎️📟📠📺🕰⏰⏲⏱🧭🎛🎚🎙📻⌛️⏳📡🔋🪫🔌💡🔦🕯🪙💷💶💴💵💸🛢🧯🪔💰💳🪪💎⚖️🪜🧰🪛🔧🧱🪤⚙️🔩🪚⛏🛠⚒🔨⛓🧲🔫💣🧨🪓🔪🗡⚔️🧿📿🔮🏺⚱️🪦⚰️🚬🛡🪬💈⚗️🔭🔬🕳🩻🩹🩺🧹🌡🧪🧫🦠🧬🩸💉💊🪠🧺🧻🚽🚰🚿🛁🛀🧼🚪🗝🔑🛎🧴🪣🧽🪒🪥🪑🛋🛏🛌🧸🪆🖼🪞🪟🎊🪅🪄🎀🎏🎈🎁🛒🛍🎉🎎🏮🎐🪩🧧✉️📩📨📫📪🪧🏷📦📤📥💌📧📬📭📮📯📜📃📄📑🧾📇🗑📅📆🗓🗒📉📈📊🗃🗳🗄📋📂📁🗂🗞📰📖📚📙📘📗📕📒📔📓🔖🧷🔗📎🖇📐📏🧮📌✏️📝🖍🖌✒️🖋🖊✂️📍🔍🔎🔏🔐🔒🔓❤️🧡💛💚💙💜🤎🖤🤍⚪️⚫️🟤🟣🔵🟢🟡🟠🔴🟥🟧🟨🟩🟦🟫⬛️⬜️❤️‍🩹❣️❤️‍🔥💔🏿🏾🏽🏼🏻💕💞💓💗💖💘💝💟☮️☦️☯️🕎🔯✡️☸️🕉☪️✝️🛐⛎♈️♉️♊️♋️♌️♍️♎️⚛️🆔♓️♒️♑️♐️♏️🉑☢️📴📳🈶🈚️🈸🈺🈷️✴️🈲🈹🈵🈴㊗️㊙️🉐💮🆚🅰️🅱️🆎🅾️🆑🆘❌⭕️🛑🚳🚯🚷♨️💢💯🚫📛⛔️🚱🔞📵🚭❗️❕❓❔‼️🔰⚜️🔱🚸⚠️〽️🔆🔅⁉️♻️✅🈯️💹❇️✳️❎🌐💠🈳🛗🅿️♿️🚾🏧💤🌀Ⓜ️🈂️🛂🛃🛄🛅🚹🚺🚼⚧🔡🔤ℹ️🔣🈁📶🎦🚮🚻🔠🆖🆗🆙🆒🆕🆓0️⃣1️⃣🔟9️⃣8️⃣7️⃣6️⃣5️⃣4️⃣3️⃣2️⃣🔢#️⃣*️⃣⏏️▶️⏸⏯⏹⏺🔽🔼◀️⏬⏫⏪⏩⏮⏭➡️⬅️⬆️⬇️↗️↘️↙️↖️↕️🔄🔂🔁🔀⤵️⤴️↩️↪️↔️🔃🎵🎶➕➖➗✖️🟰♾️🔛🔙🔚👁‍🗨®️©️™️💱💲🔝🔜〰️➰➿✔️☑️🔘🔺▫️▪️🔲🔳🔷🔶🔹🔸🔻◾️◽️◼️◻️🔈🔇🔉🔊🔔♥️♣️♠️🗯💭💬📢📣🔕♦️🃏🎴🀄️🕐🕑🕒🕓🕔🕝🕜🕛🕚🕙🕘🕗🕖🕕🕞🕟🕠🕡🕢🕣🕤🕥🕦🕧🏳️🏴🏴‍☠️🏁🚩🏳️‍🌈🏳️‍⚧️🇺🇳🇮🇸🇦🇱🇦🇼🇦🇷🇩🇿🇦🇪🇺🇸🇦🇫🇦🇿🇮🇪🇦🇲🇦🇮🇦🇴🇦🇬🇦🇩🇾🇪🇬🇧🏴󠁧󠁢󠁥󠁮󠁧󠁿🏴󠁧󠁢󠁳󠁣󠁴󠁿🇺🇬🇼🇫🇮🇩🇮🇳🇮🇷🇮🇶🇮🇹🇮🇱🏴󠁧󠁢󠁷󠁬󠁳󠁿🇺🇦🇺🇿🇺🇾🇪🇨🇪🇬🇪🇪🇸🇿🇪🇹🇪🇷🇨🇻🇬🇭🇧🇶🇳🇱🇴🇲🇦🇽🇦🇹🇦🇺🇸🇻🇬🇬🇬🇾🇰🇿🇨🇦🇮🇨🇬🇦🇨🇲🇬🇲🇰🇬🇰🇮🇬🇷🇨🇼🇨🇺🇨🇾🇬🇼🇬🇳🇰🇭🇬🇹🇬🇵🇬🇺🇰🇼🇨🇰🇬🇱🇨🇽🇬🇩🇭🇷🇨🇬🇨🇴🇰🇲🇽🇰🇨🇷🇨🇨🇨🇮🇰🇪🇰🇾🇨🇩🇸🇦🇬🇸🇼🇸🇧🇱🇸🇹🇿🇲🇵🇲🇸🇲🇸🇽🇸🇬🇸🇾🇬🇪🇯🇲🇯🇪🇬🇮🇩🇯🇸🇱🇿🇼🇨🇭🇸🇪🇸🇩🇪🇸🇸🇷🇱🇰🇸🇰🇸🇮🇸🇧🇸🇴🇱🇨🇸🇭🇻🇨🇰🇳🇷🇸🇸🇳🇸🇨🇹🇨🇹🇭🇹🇯🇨🇿🇮🇴🇹🇩🇹🇳🇨🇱🇹🇲🇹🇹🇩🇲🇩🇴🇹🇰🇹🇬🇩🇪🇩🇰🇹🇻🇹🇷🇹🇴🇳🇬🇳🇷🇳🇦🇳🇺🇳🇮🇳🇪🇳🇨🇻🇦🇵🇰🇭🇹🇧🇭🇳🇴🇳🇫🇳🇵🇳🇿🇻🇺🇧🇸🇵🇬🇧🇲🇵🇼🇵🇾🇧🇧🇵🇸🇭🇺🇫🇰🇫🇴🇵🇷🇧🇹🇫🇮🇵🇭🇫🇯🇵🇳🇧🇩🇧🇷🇫🇷🇧🇬🇧🇫🇧🇳🇧🇮🇻🇳🇧🇯🇻🇪🇵🇹🇧🇴🇧🇼🇧🇦🇵🇱🇧🇪🇵🇪🇧🇿🇧🇾🇭🇳🇲🇭🇲🇴🇲🇬🇾🇹🇲🇼🇲🇱🇲🇹🇲🇶🇲🇨🇲🇿🇲🇷🇲🇺🇲🇽🇲🇲🇫🇲🇮🇲🇲🇾🇲🇻🇲🇩🇲🇦🇲🇳🇲🇪🇲🇸🇯🇴🇱🇻🇱🇧🇱🇸🇷🇼🇱🇺🇷🇴🇱🇷🇱🇮🇱🇾🇱🇹🇷🇪🇷🇺🇻🇬🇪🇺🇰🇷🇪🇭🇭🇰🇬🇶🇨🇫🇨🇳🇹🇱🇸🇸🇦🇶🇯🇵🎌🇬🇫🇵🇫🇰🇵🇲🇵🇲🇰🇦🇸🇻🇮🇹🇫✡✳✴❇☯✂☘✉✔✘⚧✏✒♟'
```


For example, 
*  Since `EMOJI[0]` is `'😀'`, $😀 \rightarrow 0$
*  Since `EMOJI[3]` is `'😁'`, $😁 \rightarrow 3$


Use string indexing to determine the emoji that are equivalent to numbers 100, 150, and 200.

```python
emoji_100 = ...
emoji_150 = ...
emoji_200 = ...

print(emoji_100, emoji_150, emoji_200)
```

```python
grader.check("q2")
```

<hr style="border: 5px solid #346094;" />
<hr style="border: 1px solid #f4c300;" />

## Part 2: Converting Letters to Numbers

It will also be important that we can convert numerical representations back into their characters, as it is often easier to transmit characters instead of raw numerical data. To convert numbers into their equivalent character, we can use the `.find()` method.

The `A.find(B)` method looks for string `B` inside of string `A`, and if it finds it, returns the index of string `A` where it was found.

For example,
* `'😀😃😄😁😆🥹'.find('😆')` would return 4, since the string `'😆'` is found at index 4 inside of the string `'😀😃😄😁😆🥹`.
* `LETTERS.find('C')` would return 2, since the string `'C'` is found at index 2 inside of the string `LETTERS`.

### Question 3

Use the `.find()` method with the `BIGLETTERS` alphabet to decode the following message, provided as a sequence of numbers:

`0, 28, 26, 29, 34, 26, 52`

When you've determined the plaintext word, assign it as a string to the variable `park`.

```python
# Write your code here to determine the corresponding characters in BIGLETTERS
# Use a print statement to help see the results of any intermediate calculations
...

park = ...
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

### Question 4

What happens if you try to use the `A.find(B)` method, but Python can't find the substring B instead of the string A? Write code where this happens, and then describe what you observe. 

Is there an error message? If yes, what does the error message look like? If no, what happens?

```python
# Write your code here to investigate
...
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 5

There is another Python method `.index()` that works very similarly to `.find()`, but operates differently when it can't find a substring A in the larger string B. 

**Investigate the difference in the code cell below, and explain why you might want to choose `.index()` instead of `.find()`.**

```python
# Explore the difference between index and find
...
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #346094;" />
<hr style="border: 1px solid #f4c300;" />


## Part 3: The Caeser Cipher 


### The Modulo Operator

The modulo operator, `%`, in Python computes the remainder after division. Running the command `7 % 3` is equivalent to solving the mathematical statement: $7 \equiv x \pmod 3$, where the `%` operation yields the smallest positive value for $x$.

Evaluate by hand, then check using Python:
* `7 % 26`
* `26 % 26`
* `27 % 26`
* `30 % 26`
* `52 % 26`

```python
# Use this code cell to explore how the % operator works
...
```

### Question 6

The Caesar cipher is an **additive cipher**, which means that to encipher a character from plaintext to ciphertext, you **add an integer value**, which we call the *key*, to the numerical representation of the plaintext character to determine the numerical representation of the ciphertext character.


For example:
* If our `key = 5`, then plaintext `b`, which numerically is `1`, becomes `6`, which corresponds to `G`
* If our `key = 11`, then plaintext `h`, which numerically is `7`, becomes `18`, which corresponds to `S`

Fill in the sections of `...` so the code in the cell below computes the **numerical values** of the ciphertext characters that correspond to the plaintext word  `MATH` when using the key value of `3`.

```python
key = ...

M_ciphertext_number = LETTERS.index(...) + ...
A_ciphertext_number = ...
T_ciphertext_number = ...
H_ciphertext_number = ...

print(M_ciphertext_number, A_ciphertext_number, T_ciphertext_number, H_ciphertext_number)
```

```python
grader.check("q6")
```

### Question 7

Now, fill in the sections of `...` so the code in the cell below determines the **ciphertext characters** that correspond to the plaintext word `MATH` when using the key value of 5. You should be able to encipher each individual character using a single line of code.

```python
key = ...

M_enciphered = LETTERS[ LETTERS.index( ... ) + ... ]
A_enciphered = LETTERS[ ... ]
T_enciphered = ...
H_enciphered = ...

print(M_enciphered, A_enciphered, T_enciphered, H_enciphered)
```

```python
grader.check("q7")
```

You may have already noticed that there are combinations of key values and plaintext numerical representations that sum to a value larger than 25, which means we obtain a ciphertext numerical value that doesn't have a corresponding letter. 

For example,

If `key = 15` and the plaintext character is `q`, then the ciphertext numerical value is `15 + 16 = 31`. There is no character in `LETTERS` with a corresponding value of `31`.

However, there is an equivalent numerical value between 0 and 25 to 31 that we can use instead. The modulo operator, `%` can compute the corresponding value!

For example, `31 % 26` computes 5, which is equivalent to `F`.

### Question 8

Fill in the `...` sections below so the code determines the **ciphertext characters** that correspond to the plaintext word `DOG` when using the **key value of 23**. Your code should work even if the result of adding the key to the plaintext value results in a number that's too large for the alphabet. You can copy/paste some of your code from the previous question to get started.

```python
key = ...

D_enciphered = ...
O_enciphered = ...
G_enciphered = ...

print(D_enciphered, O_enciphered, G_enciphered)
```

```python
grader.check("q8")
```

## Question 9
The Hawaiian alphabet (*in Hawaiian: ka pīʻāpā Hawaiʻi*) only has 13 characters: A, E, I, O, U, H, K, L, M, N, P, W, and '. We will define this set of characters as `HAWAIIAN` in Python using the cell of code below.

```python
HAWAIIAN = "AEIOUHKLMNPW'"
```

Write Python code below that **deciphers** the each character of the ciphertext message `HLIWI` using a key value of 15, and the alphabet `HAWAIIAN`

**HINT:** Think about how you will need to change your code from the previous questions to account for the fact that you are **deciphering** AND are using an alphabet with **13 characters**, not the 26 found in `LETTERS`.

```python
key = ...

H_deciphered = ...
L_deciphered = ...
I_deciphered = ...
W_deciphered = ...
I_deciphered = ...

print(H_deciphered, L_deciphered, I_deciphered, W_deciphered, I_deciphered)
```

```python
grader.check("q9")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: 1)
```
>>> character_31 == 'f'
True
```
failure_message: ❌ You did not correctly determine the character for 31
success_message: ✅ You correctly determined the character for 31
```
>>> character_44 == 's'
True
```
failure_message: ❌ You did not correctly determine the character for 44
success_message: ✅ You correctly determined the character for 44
```
>>> character_47 == 'v'
True
```
failure_message: ❌ You did not correctly determine the character for 47
success_message: ✅ You correctly determined the character for 47
```
>>> character_53 == '@'
True
```
failure_message: ❌ You did not correctly determine the character for 53
success_message: ✅ You correctly determined the character for 53


Tests for q2 (points: 1)
```
>>> emoji_100
'💫'
```
failure_message: ❌ You did not correctly determine the emoji for 100
success_message: ✅ You correctly determined the emoji for 100
```
>>> emoji_150
'🫰'
```
failure_message: ❌ You did not correctly determine the emoji for 150
success_message: ✅ You correctly determined the emoji for 150
```
>>> emoji_200
'👶'
```
failure_message: ❌ You did not correctly determine the emoji for 200
success_message: ✅ You correctly determined the emoji for 200


Tests for q3 (points: 1)
```
>>> hash(park) == 2534782857975858343
False
```
failure_message: ❌ You did not correctly determine the plaintext word
success_message: ✅ You correctly determined the plaintext word!


Tests for q6 (points: 1)
```
>>> M_ciphertext_number
15
```
failure_message: ❌ You did not correctly determine ciphertext value for plaintext character M
success_message: ✅ You correctly determined the ciphertext value for plaintext character M!
```
>>> A_ciphertext_number
3
```
failure_message: ❌ You did not correctly determine ciphertext value for plaintext character A
success_message: ✅ You correctly determined the ciphertext value for plaintext character A!
```
>>> T_ciphertext_number
22
```
failure_message: ❌ You did not correctly determine ciphertext value for plaintext character T
success_message: ✅ You correctly determined the ciphertext value for plaintext character T!
```
>>> H_ciphertext_number
10
```
failure_message: ❌ You did not correctly determine ciphertext value for plaintext character H
success_message: ✅ You correctly determined the ciphertext value for plaintext character H!


Tests for q7 (points: 1)
```
>>> M_enciphered
'R'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character M
success_message: ✅ You correctly determined the ciphertext character for plaintext character M!
```
>>> A_enciphered
'F'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character A
success_message: ✅ You correctly determined the ciphertext character for plaintext character A!
```
>>> T_enciphered
'Y'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character T
success_message: ✅ You correctly determined the ciphertext character for plaintext character T!
```
>>> H_enciphered
'M'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character H
success_message: ✅ You correctly determined the ciphertext character for plaintext character H!


Tests for q8 (points: 1)
```
>>> D_enciphered
'A'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character D
success_message: ✅ You correctly determined the ciphertext character for plaintext character D!
```
>>> O_enciphered
'L'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character O
success_message: ✅ You correctly determined the ciphertext character for plaintext character O!
```
>>> G_enciphered
'D'
```
failure_message: ❌ You did not correctly determine ciphertext character for plaintext character G
success_message: ✅ You correctly determined the ciphertext character for plaintext character G!


Tests for q9 (points: 1)
```
>>> H_deciphered != 'L'
True
```
failure_message: ❌ You are still adding the key, even though you should be deciphering
```
>>> H_deciphered
'O'
```
failure_message: ❌ You did not correctly determine plaintext character for ciphertext character H
success_message: ✅ You correctly determined the plaintext character for ciphertext character H!
```
>>> L_deciphered
'H'
```
failure_message: ❌ You did not correctly determine plaintext character for ciphertext character H
success_message: ✅ You correctly determined the plaintext character for ciphertext character H!
```
>>> I_deciphered
'A'
```
failure_message: ❌ You did not correctly determine plaintext character for ciphertext character H
success_message: ✅ You correctly determined the plaintext character for ciphertext character H!
```
>>> W_deciphered
'N'
```
failure_message: ❌ You did not correctly determine plaintext character for ciphertext character H
success_message: ✅ You correctly determined the plaintext character for ciphertext character H!
